# Lab 01: Basic Python

## Exercise 1: Perfect Numbers

Most of the lab notebooks you'll be working on for this class will come with a fair bit of skeleton code --- i.e., stubbed out classes and functions that you need to complete or modify to get working correctly.

In the cell below, for instance, you'll find a stubbed out function named `is_perfect`, which should return `True` if the number passed to it is a "perfect" number, and `False` otherwise.

A perfect number is a postive integer whose value is equal to the sum of its proper divisors (i.e., its factors excluding the number itself). 6 is the first perfect number, as its divisors 1, 2, and 3 add up to 6.

Fill in your own implementation of the function below:

<!-- grade_id: is_perfect, points: 0, grade: false, solution: true -->
```python
def is_perfect(n):
    if n == 0 or n < 0 or type(n) != int:
        return False
    else:
        divisors = []
        for x in range(1,n):
            if n % x == 0:
                divisors.append(x)

        sum = 0
        for i in divisors:
            sum += i
        return sum == n    
```

Each exercise will also be accompanied by one or more *unit test* cells, each of which is meant to test some aspect of your implementation. When you run the unit test cell(s) after evaluating your implementation, you'll either find errors reported, which should help you identify what you need to fix, or they will complete silently, which means you've passed the test(s).

**It's important that you ensure your implementation and test cell(s) actually run to completion before moving on** --- there's a big difference between a cell not producing an error and not completing! (A "`In [*]`" marker next to the cell means that it's still being evaluated by the interpreter.)

We will often run additional *hidden tests* upon submission of your notebook to guard against hardcoded solutions.  While you won't see the specific test cases, you can resubmit your work as many times as you wish (before the deadline) to ensure they pass.

<!-- grade_id: is_perfect_test, points: 3, grade: true, solution: false -->
```python
# (3 points)
import unittest
tc = unittest.TestCase()

for n in (6, 28, 496):
    tc.assertTrue(is_perfect(n), '{} should be perfect'.format(n))

for n in (1, 2, 3, 4, 5, 10, 20):
    tc.assertFalse(is_perfect(n), '{} should not be perfect'.format(n))

for n in range(30, 450):
    tc.assertFalse(is_perfect(n), '{} should not be perfect'.format(n))
```

## Exercise 2: Multiples of 3 and 5

If we list all the natural numbers below 10 that are multiples of 3 or 5, we get 3, 5, 6 and 9. The sum of these multiples is 23.

Complete the following function, which finds the sum of all the multiples of 3 or 5 below the argument `n`.

<!-- grade_id: multiples_of_3_and_5, points: 0, grade: false, solution: true -->
```python
def multiples_of_3_and_5(n):
    mults = []
    for i in range(1, n):
        if i % 3 == 0 or i % 5 == 0:
            mults.append(i)
    sum = 0
    for i in mults:
        sum += i
        
    return sum
```

<!-- grade_id: multiples_of_3_and_5_test, points: 3, grade: true, solution: false -->
```python
# (3 points)
import unittest
tc = unittest.TestCase()
tc.assertEqual(multiples_of_3_and_5(10), 23)
tc.assertEqual(multiples_of_3_and_5(500), 57918)
tc.assertEqual(multiples_of_3_and_5(1000), 233168)
```

## Exercise 3: Integer Right Triangles

Given a perimeter of 60, we can find two right triangles with integral length sides: [(10, 24, 26), (15, 20, 25)]. Complete the following function, which takes an integer `p` and returns the number of unique integer right triangles with perimeter `p`.

Note that your solution should take care to limit the number of triangles it tests --- **your function must complete in under 3 seconds for all values of `p` used in the test cells below to earn credit.**

<!-- grade_id: integer_right_triangles, points: 0, grade: false, solution: true -->
```python
def integer_right_triangles(p):
    #a + b + c = p and a^2 + b^2 = c^2
    # b = p - c - a
    # b = sqrt(c^2 - a^2)
    #with some mathematical manipulation & simplification, 

    #If p is odd, return 0 because there won't be any 
    if p % 2 != 0:
        return 0
    else:
        count = 0
        for a in range(1,p):
            for b in range(a, p-a):
                if(a*a + b*b == (p - (a+b))*(p - (a+b))):
                     count += 1
    return count 
```

<!-- grade_id: integer_right_triangles_tests_1, points: 1, grade: true, solution: false -->
```python
# (2 points)
import unittest
tc = unittest.TestCase()
tc.assertEqual(integer_right_triangles(60), 2)
tc.assertEqual(integer_right_triangles(100), 0)
tc.assertEqual(integer_right_triangles(180), 3)
```

## Exercise 4: Simple ASCII Art

For this next exercise, you'll need to complete the function `gen_pattern`, which, when called with a string of length $\ge$ 1, will print an ASCII art pattern of concentric diamonds using those characters. The following are examples of patterns printed by the function (note the newline at the end of the last line!):

    > gen_pattern('X')
    
    X
    
    > gen_pattern('XY')
    
    ..Y..
    Y.X.Y
    ..Y..
    
    > gen_pattern('WXYZ')
    
    ......Z......
    ....Z.Y.Z....
    ..Z.Y.X.Y.Z..
    Z.Y.X.W.X.Y.Z
    ..Z.Y.X.Y.Z..
    ....Z.Y.Z....
    ......Z......
    
You ought to find the string [`join`](https://docs.python.org/3.6/library/stdtypes.html#str.join) and [`center`](https://docs.python.org/3.6/library/stdtypes.html#str.center) methods helpful in your implementation. They are demonstrated here:

    > '*'.join('abcde')
    
    'a*b*c*d*e'
    
    > 'hello'.center(11, '*')
    
    '***hello***'
    
Complete the `gen_pattern` function, below:

<!-- grade_id: gen_pattern, points: 0, grade: false, solution: true -->
```python
def gen_pattern(chars):
    char_len = len(chars)

# let example be 'abc'
#for the centering to find out how many terms in a line: (char_len * 4) - 3
    center_val = (char_len*4) - 3
  

    
    for i in range(char_len - 1, -1, -1): #for the first half of the pyramid (we want 'i' to iterate as 2,1,0)

        # print_term for the backwards dir

        print_term_b = chars[char_len:i:-1] + chars[i:char_len]
        #add the '.'

        dot_add_b = '.'.join(print_term_b)
  

        # utilize center

        centered_term_b = dot_add_b.center(center_val, '.')


        print(centered_term_b) #printing the first half of the pyramid
        
    #for loop for the bottom half now
                                                                                           
    for i in range(1, char_len): #forward terms #want 'i' to iterate as 1,2
        # print_term for the forward dir. 

        print_term_f = chars[char_len:i:-1] + chars[i:char_len]
        
        #add the '.'

        dot_add_f = '.'.join(print_term_f)
  

        # utilize center

        centered_term_f = dot_add_f.center(center_val, '.')

        print(centered_term_f)


```

```python
# (1 point) output:
# @

gen_pattern('@')
```

```python
# (1 point) output:
# ..%..
# %.@.%
# ..%..

gen_pattern('@%')
```

```python
# (1 point) output:
# ....C....
# ..C.B.C..
# C.B.A.B.C
# ..C.B.C..
# ....C....

gen_pattern('ABC')
```

```python
# (1 point) output:
# ........#........
# ......#.#.#......
# ....#.#.#.#.#....
# ..#.#.#.#.#.#.#..
# #.#.#.#.#.#.#.#.#
# ..#.#.#.#.#.#.#..
# ....#.#.#.#.#....
# ......#.#.#......
# ........#........

gen_pattern('#####')
```

<!-- grade_id: gen_pattern_tests_2, points: 2, grade: true, solution: false -->
```python
# (2 points) output:
# ..............................p..............................
# ............................p.o.p............................
# ..........................p.o.n.o.p..........................
# ........................p.o.n.m.n.o.p........................
# ......................p.o.n.m.l.m.n.o.p......................
# ....................p.o.n.m.l.k.l.m.n.o.p....................
# ..................p.o.n.m.l.k.j.k.l.m.n.o.p..................
# ................p.o.n.m.l.k.j.i.j.k.l.m.n.o.p................
# ..............p.o.n.m.l.k.j.i.h.i.j.k.l.m.n.o.p..............
# ............p.o.n.m.l.k.j.i.h.g.h.i.j.k.l.m.n.o.p............
# ..........p.o.n.m.l.k.j.i.h.g.f.g.h.i.j.k.l.m.n.o.p..........
# ........p.o.n.m.l.k.j.i.h.g.f.e.f.g.h.i.j.k.l.m.n.o.p........
# ......p.o.n.m.l.k.j.i.h.g.f.e.d.e.f.g.h.i.j.k.l.m.n.o.p......
# ....p.o.n.m.l.k.j.i.h.g.f.e.d.c.d.e.f.g.h.i.j.k.l.m.n.o.p....
# ..p.o.n.m.l.k.j.i.h.g.f.e.d.c.b.c.d.e.f.g.h.i.j.k.l.m.n.o.p..
# p.o.n.m.l.k.j.i.h.g.f.e.d.c.b.a.b.c.d.e.f.g.h.i.j.k.l.m.n.o.p
# ..p.o.n.m.l.k.j.i.h.g.f.e.d.c.b.c.d.e.f.g.h.i.j.k.l.m.n.o.p..
# ....p.o.n.m.l.k.j.i.h.g.f.e.d.c.d.e.f.g.h.i.j.k.l.m.n.o.p....
# ......p.o.n.m.l.k.j.i.h.g.f.e.d.e.f.g.h.i.j.k.l.m.n.o.p......
# ........p.o.n.m.l.k.j.i.h.g.f.e.f.g.h.i.j.k.l.m.n.o.p........
# ..........p.o.n.m.l.k.j.i.h.g.f.g.h.i.j.k.l.m.n.o.p..........
# ............p.o.n.m.l.k.j.i.h.g.h.i.j.k.l.m.n.o.p............
# ..............p.o.n.m.l.k.j.i.h.i.j.k.l.m.n.o.p..............
# ................p.o.n.m.l.k.j.i.j.k.l.m.n.o.p................
# ..................p.o.n.m.l.k.j.k.l.m.n.o.p..................
# ....................p.o.n.m.l.k.l.m.n.o.p....................
# ......................p.o.n.m.l.m.n.o.p......................
# ........................p.o.n.m.n.o.p........................
# ..........................p.o.n.o.p..........................
# ............................p.o.p............................
# ..............................p..............................

gen_pattern('abcdefghijklmnop')
```